<a href="https://colab.research.google.com/github/yy-0715/DSSS_homework2/blob/main/0_setup.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Fill in your github username and idm credentials below:

```python
github_username = "yy-0715"
idm_credentials = "oc49uxyc"
```

---

# **Applied Data Science in Medicine & Psychology** - Assignment 0
## Setup
---

This first assignment of the "Applied Data Science in Medicine & Psychology" course introduces the tools used for the programming tasks and gives instructions about their setup.

### **General Setup and Submission**
1. Save an editable copy of this notebook to your Google Drive ``File`` $\rightarrow$ ``Save a copy in Drive`` and remove the automatically appended "Copy of" text. **Do not change the original file name**.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** (otherwise this causes an error in your implementation) or *YOUR ANSWER HERE* with your solution.
3. If you take a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When you continue working on the assignment, do this in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your repository.

### **Assignment Description**

#### **Assignment Goals**
You are going to learn the following things in this assignment:
1. How to setup your Google Colaboratory and GitHub accounts
2. How to access and submit assignments

#### **Assignment Point Value**
This assignment is worth a total (up to) **1** point. If your code passes all of the automatic tests, you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

1. Change the code in the next cell to indicate that you have gone through the exercise introduction in StudOn and set up your Google Colaboratory and GitHub account. You should change each variable from ``False`` to ``True`` to show which tasks you have accomplished (by self-report).

<!-- grade_id: task_1, points: 0, grade: false, solution: true -->
```python
# I have read through the introduction about Tools, Grading and Setup
read_introduction = True

# I have a Google account and can access Google Colaboratory under my own
# account
access_colaboratory = True

# I have a GitHub account
github_account = True
```

By running the following cell you can check for errors in your implementation. Make sure that no "AssertationError" occurs.

<!-- grade_id: task_1_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert read_introduction == True, "The variable read_introduction is not set to True"
assert access_colaboratory == True, "The variable access_colaboratory is not set to True"
assert github_account == True, "The variable github_account is not set to True"
```

2. Also update the ``github_username`` and ``idm_credentials`` accordingly. Keep the **'** around **'**username**'** and **'**ab12abcd**'**.

```python
# enter your github username
github_username = 'yy-0715'

# enter your FAU idM-credentials
idm_credentials = 'oc49uxyc'
```

### **Submission**
After you finished completing the code above, you need to submit your assignment. Submitting the assignment entails syncing your copy of the notebook with your private GitHub assignment repository (``digikolleg-data-science-psychology/0-setup-<username>``):
* Check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning.
* Save your notebook, click ``File`` $\rightarrow$ ``Save``
* Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
* Select ``digikolleg-data-science-psychology/0-setup-<username>`` from the repository drop-down menu
* Type ``0_setup.ipynb`` into the file path text box
* Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
* Make sure that the "Include a link to Colaboratory" box is checked
* Press "Ok"
* A new tab or window should open up, displaying your submitted notebook within your GitHub repository

### **Updating your submission**
Prior to the assignment deadline you can re-submit your assignment as often as you would like. This will update the copy of your assignment file in your assignment's repository. Update the assignment from the notebook that you saved in your Google Drive (access from your Google Colab account) or your own repository.